Before you turn this assignment in, make sure everything runs as expected. To do that select in the menubar Runtime $\rightarrow$ Restart and run all. This will clear all your local variables and runs your notebook from the beginning.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your GitHub username and your idm credentials below:

```python
github_username = ""
idm_credentials = ""
```

---

# **Applied Data Science in Medicine & Psychology** - Assignment 1
---
***Deadline: dd.mm.2023, hh:mm:ss TODO***

This first assignment of the "Applied Data Science in Medicine & Psychology" course introduces the tools used for the programming tasks and gives instructions about their setup.

### **Tools**
We will use the following tools in this course:
* [Google Colaboratory](https://colab.research.google.com/) (Colab):
a system which is included as part of all Google accounts for writing and running code. This resource provides an easy means of organizing notes, code, and graphics in a single cohesive format (called "notebook"). In addition to being a generally useful tool, Colaboratory has the benefit that all code will run on Google Cloud servers (rather than your personal computer). This substantially reduces setup costs and difficulties. 
* [Github](https://github.com/): a server used to manage and share data and code. It provides an easy way of managing multiple versions of data and code that may be easily shared and tracked. Through GitHub you will access the assignment notebooks.
* Github Classroom: all assignments are organized in a GitHub classroom. After accepting the assignment a repository with all the corresponding files is generated which is only accessible for you and your course instructor. A message should inform you where it is located and it is also accessible from your GitHub account. The repository serves as a framework for solving the assignment. After working with the provided material you will need to submit it for grading.

### **Assignment Grading**
All course assignments will be assigned a number of points. In order to gain those points you have to follow the instructions in the notebooks and complete the given code. After submission, your code will be evaluated automatically as follows:
* You can submit and modify your assignment up until the assignment deadline
* When the assignment deadline is reached, a series of tests (test cases) will be automatically applied to your assignment:
  * each test case will be associated with a specific point value
  * If the test case *passes* you will get the full number of points for this test case
  * if the test case *fails* you will get 0 points for this case
  * the grade you receive for the assignment will be calculated by adding together all of the points from the *passed* test cases
* Different assignments are worth different amounts of points, according to ther difficulty and complexity
* If you miss the assignment deadline, you will get 0 points for this assignment

### **Setup**
1. Create a free [Github account](https://github.com/). If you already have an account, you may use it for this course.
2. Create a [Google account](https://accounts.google.com/signup/v2/webcreateaccount?flowName=GlifWebSignIn&flowEntry=SignUp). If you already have an account, you may use it for this course. Make sure you can sign into Google Colaboratory using this [link](https://colab.research.google.com/).
3. Ensure that Google Colaboratory has access to your private repositories on GitHub. (By default, your assignment repository will be accessible only to you and the course instructor). To grant access:
  * Navigate to [https://colab.research.google.com/github/](https://colab.research.google.com/github/)
  * Make sure the "Include private repositories" box is checked
  * In the popup window, sign into your Github account and authorize Colab to read private files
4. In order to generate your own assignment repository and be able to submit your assignment, you need to accept the assignment using this [link]()**TODO**.  
5. Click on the blue "Open in Colab" button on the top of the page. A new tab or window should open in your browser, displaying this notebook in Google Colaboratory.
6. **Continue the setup of your assignment whithin Google Colaboratory with the following point**
7. To be able to edit and submit the assignment you need to save an editable copy of your notebook. Click ``File > Save a copy in Drive`` and then save your copy as ``ads_assignment_1`` (i.e. remove the "Copy of" text that gets automatically appended to the beginning of the filename).

### **Jupyter (Colaboratory) Notebooks**
A Jupyter/Colaboratory Notebook is an interactive tool that allows you to take notes, execute code, create figures, and easily share your work with other people. It always has the file extension .ipynb which stands for **i**nteractive **Py**thon **n**ote**b**ook. Notebooks comprise a series of cells. Each cell (looks like a rectangular block of text) in a notebook contains one of two basic categories of content:
1. **Markdown cells**: contains formatted text. We use this type of cell to give explanations and instructions about the assignment tasks.
2. **Code cells**: contains computer code (in our case Python code) which is a set of instructions that tell the notebook what to do when the cell is *run*. To execute/run the code in a cell hold the ``shift`` key and press ``enter/return`` or press the start button in the upper left corner of the cell. 

### **Assignment Description**

#### **Assignment goals**
You are going to learn the following things in this assignment:
1. How to setup your Google Colaboratory and GitHub accounts
2. How to access and submit assignments

#### **Assignment point value**
This assignment is worth a total (up to) **1** point. If your code passes all of the automatic tests, you will receive full credit.

### **Tasks**
Make sure that you now work on your *copied* version of this assignment's notebook.

Change the code in the next cell to indicate that you have gone through the introduction and set up your Google Colaboratory and GitHub account. You should change each variable from ``False`` to ``True`` to show which tasks you have accomplished (by self-report). By running the cell you can check for errors in your implementation. Make sure that no "AssertationError" occurs.

<!-- grade_id: variables_test, points: 1, grade: true, solution: false -->
```python
# I have read through the introduction about Tools, Grading and Setup
read_introduction = False

# I have a Google account and can access Google Colaboratory under my own 
# account
access_colaboratory = False

# I have a GitHub account
github_account = False

# tests for checking your implementation
assert read_introduction == True, "The variable read_introduction is not set to True"
assert access_colaboratory == True, "The variable access_colaboratory is not set to True"
assert github_account == True, "The variable github_account is not set to True"
```

Also update the ``github_username`` and ``idm_credentials`` accordingly. Keep the **'** around **'**username**'** and **'**ab12abcd**'**.

<!-- grade_id: user_info_test, points: 1, grade: true, solution: false -->
```python
# enter your github username
github_username = 'username'

# enter your FAU idM-credentials
idm_credentials = 'ab12abcd'

assert isinstance(github_username, str), "The variable github_username is not of type string"
assert isinstance(idm_credentials, str), "The variable idm_credentials is not of type string"
```

### **Submission**
After you finished completing the code above, you need to submit your assignment. Submitting the assignment entails syncing your copy of the notebook with your private GitHub assignment repository (``Applied Data Science in Medicine & Psychology/ads_assignment_1-<username>``):
1. Save your notebook, click ``File > Save``
2. From within Google Colaboratory, click ``File > save a copy in Github``
3. Select ``Applied Data Science in Medicine & Psychology/ads_assignment_1-<username>`` from the repository drop-down menu
4. Type ``ads_assignment_1.ipynb`` into the File path text box
5. Optionally, you may enter a description into the "Commit message" text box (like: "solved task 1" or "ready for grading")
6. Make sure that the "Include a link to Colaboratory" box is checked
7. Press "Ok"

A new tab or window should open up, displaying your submitted notebook within your GitHub repository.

### **Updating your submission**
Prior to the assignment deadline you can re-submit your assignment as often as you would like. This will update the copy of your assignment file in your assignment's repository.

#### **Re-starting your submission from scratch**
To restart your submission from scratch (from the original template notebook for this assignment), simply re-open the assignment's notebook in Colaboratory using this [link](***TODO***) ***TODO***.

#### **Modifying your prior submission(s)**
As an alternative to starting from scratch, you can also use your existing assignment as a jumping off point for further edits. The easiest way to do this is to open your Google Colaboratory account ([link](https://colab.research.google.com/)) and select your assignment's notebook using the dialogue box. It should be in the Recent tab. Or, if you've been working on other files, you can navigate to your notebook within the Google Drive (all Colaboratory notebooks are stored on your Google Drive space) or GitHub tabs. After saving your changes, simply follow the submission instructions again to re-copy your assignment to GitHub.