Fill in your github username and idm credentials below:

```python
github_username = ""
idm_credentials = ""
```

---

# **Applied Data Science in Medicine & Psychology - Assignment 3**
---
***Deadline: dd.mm.2023, hh:mm:ss TODO***
### **General setup and submission**
1. Save an editable copy of this notebook to your Google Drive ``File > Save a copy in Drive`` and remove the automatically appended "Copy of" text. Do not change the file name.
2. Work on the assignment in the *copied* version of the notebook. Make sure that you **replace** all parts of the tasks indicated with *YOUR CODE HERE* and **raise NotImplementedError()** or *YOUR ANSWER HERE* with your solution.
3. If you do a break during the implementation, you can save the current status of your work in your repository following the submission instructions below. When continuing the assignment **do not** reaccept the assignment, this would overwrite your work done so far. Continue the assignment in the version of the notebook that you saved in your Google Drive (access from your Google Colab account) or your own repository.
4. After finishing the assignment submit your assignment as follows:
    * check that everything runs as expected. To do that select in the menubar ``Runtime`` $\rightarrow$ ``Restart and run all``. This will clear all your local variables and runs your notebook from the beginning
    * Save your notebook, click ``File`` $\rightarrow$ ``Save``
    * Click ``File`` $\rightarrow$ ``Save a copy in GitHub``
    * Select ``Applied Data Science in Medicine & Psychology/ads_assignment_3-<username>`` from the Repository drop-down menu 
    * Type ``ads_assignment_3.ipynb`` into the File path text box
    * Optionally, you may enter a description into the "Commit message" text box (like: "solved task 1" or "ready for grading")
    * Make sure that the "Include a link to Colaboratory" box is checked
    * Press "Ok"
    * A new tab or window should open up, displaying your submitted notebook within your GitHub repository

Before the assignment deadline, you can re-submit your assignment as often as you would like. This will update the copy of your assignment file in your assignment's repository.

### **Assignment Description**
In this assignment, we cover the second part of Python basics. You will learn essential programming building blocks such as functions, data types and structures, control flow elements, and data type casting which are used to store, manipulate, and control data in a program. Understanding these concepts is crucial for writing efficient and effective code.

#### **Assignment goals**
You are going to learn the following things in this assignment:
1. how to implement and use functions
2. how to create variables of different data types
3. how to set up and work with dictionaries and lists
4. how to control the program flow with if-statements and loops
5. how to cast variables to different datatypes

#### **Assignment point value**
This assignment is worth a total (up to) **12** points. If your code passes all of the manual and automatic tests you will receive full credit.

### **Tasks**
Make sure that you now work on your *copied* version of this assignment's notebook.

#### 1. Functions
Functions are a set of instructions that perform a specific task within a program. They allow developers to modularize code and reuse it throughout the program. Functions can take arguments (inputs) and return values (outputs):
```
def <function_name>(<arg_1>, ..., <arg_N>):
    <instruction_1>
    ...
    <instruction_N>
    return <value>
```

##### 1.1 Create a function called ``get_datatype`` that takes a variable called ``var`` as argument and returns the type of the variable:
**Hint:** check the lecture slides for a suitable function. Test your solution with different datatypes.

<!-- grade_id: function_datatype, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: function_datatype_test, points: 0.5, grade: true, solution: false -->
```python
# run tests to check your implementation
assert get_datatype(5) == int, "check if you return the data type itself and not an string representation or similar"
assert get_datatype("Hello, world!") == str, "check if you return the data type itself and not an string representation or similar"
```

#### 2. Datatypes - Dictionary
A dictionary is a type of data structure that stores key-value pairs. Each key in the dictionary is associated with a value:
```
dict_1 = {
    key_1: value_1,
    key_2: value_2
}
```
While dictionary keys must be unique, dictionary values can be duplicated. To access the value associated with a particular key, you can use square brackets:
```
dict_1[key_1]
```
##### 2.1 Create a dictionary called ``patient`` with the following key-value pairs:
| key | value | datatype |
| :- | :- | :- |
|name|Jane|string|
|height|167.3|float|
|vegetarian|True|boolean|
|allergies|pollen, peanuts|list|
|blood_test_results|None|Null|

<!-- grade_id: dict_patient, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: dict_patient_test, points: 1, grade: true, solution: false -->
```python
# run tests to check your implementation
assert patient
assert patient["name"] == "Jane", "check the lecture slides how to create values of certain datatypes. Make sure that you use the correct datatype for the keys"
```

##### 2.2 Add a key-value pair to the ``patient`` dictionary:
| key | value | datatype |
| :- | :- | :- |
|age|24|int|

**Hint:** you can add a key-value pair by using the accessing strategy described above and assigning a value to the specified key.

<!-- grade_id: dict_add, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
print(patient)
```

<!-- grade_id: dict_add_test, points: 0.5, grade: true, solution: false -->
```python
# run tests to check your implementation
assert patient["age"], "check if you added the key-value pair correctly"
```

##### 2.3 Create another dictionary called ``blood_test_results`` and assign it to the ``blood_test_results`` key of ``patient``:
| key | value | datatype |
| :- | :- | :- |
|hemoglobin|13.6|float|
|potassium|4.5|float|

<!-- grade_id: dict_dict, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
print(patient)
```

<!-- grade_id: dict_dict_test, points: 1, grade: true, solution: false -->
```python
# run tests to check your implementation
assert blood_test_results
assert len(blood_test_results) == 2, "check if you added both elements to the dictionary"
```

#### 3. Datatypes - List
A Python list is a structure that allows you to store a collection of elements of any data type in a single variable, like
```
python_list = [12, "hello", 23.5, 25]
```
Lists are provided with the following functionalities:

| functionality | key word|explanation | example | return |
| :- | :- | :- | :- | :- |
|Accessing||access individual elements by their index, <br /> index starts from 0|``python_list[0]``|``12``|
|Slicing||create a sub-list|``python_list[1:3]``|``["hello", 23.5]``|
|Inserting|``insert``|insert a new element at a specific position|``python_list.insert(2, "new")`` | ``[12, "hello", "new", 23.5, 25]``
|Deleting|``remove``|remove an element|``python_list.remove("hello")``|``[12, 23.5, 25]``|
|Appending|``append``|add new elements to the end|``python_list.append([26, 27])``|``[12, "hello", 23.5, 25, [26, 27]]``|
|Extending|``extend``|add new elements to the end|``python_list.extend([26, 27])``|``[12, "hello", 23.5, 25, 26, 27]``|
|Membership|``in``|check if a particular element is present|``25 in python_list`` <br /> ``56 in python_list``|``True`` <br /> ``False``

##### 3.1 Create a list called ``my_list`` with the elements ``1, 2, "hello", 4, True``:

<!-- grade_id: list, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: list_test, points: 0.5, grade: true, solution: false -->
```python
# run tests to check your implementation
assert my_list
assert len(my_list) == 5, "check if you created the list with all the elements"
assert my_list[2] == "hello", "check if all elements are in the right order"
```

##### 3.2 Create a variable ``result`` which adds the 2nd and 4th element of ``my_list`` using their index:

<!-- grade_id: list_index, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: list_index_test, points: 0.5, grade: true, solution: false -->
```python
# run tests to check your implementation
assert result
assert result == 6
```

##### 3.3 Create a list ``sub_list`` which contains the 3rd, 4th, and 5th element of ``my_list`` using slicing:

<!-- grade_id: list_slicing, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: list_slicing_test, points: 0.5, grade: true, solution: false -->
```python
# run tests to check your implementation
assert sub_list
assert len(sub_list) == 3, "check the length of your sub_list and your indices. Did you consider how the last element is handled in slicing?"
```

##### 3.4 Insert at the 4th position of ``new_list`` the string ``world``. Delete ``4`` out of ``new_list``:

<!-- grade_id: list_ins_del, points: 0, grade: false, solution: true -->
```python
new_list = my_list.copy()

# YOUR CODE HERE
raise NotImplementedError()

print(new_list)
```

<!-- grade_id: list_ins_del_test, points: 0.5, grade: true, solution: false -->
```python
# run tests to check your implementation
assert len(new_list) == 5, "check if you performed both operations (inserting, deleting)"
assert new_list[3] == "world", "check if you used the correct index for inserting"
```

##### 3.5 Append ``tmp_list`` to ``new_list``:

<!-- grade_id: list_append, points: 0, grade: false, solution: true -->
```python
tmp_list = [False, "hi", 3]
new_list = my_list.copy()

# YOUR CODE HERE
raise NotImplementedError()

print(new_list)
print(f"length of new_list: {len(new_list)}")
```

<!-- grade_id: list_append_test, points: 0.5, grade: true, solution: false -->
```python
# run tests to check your implementation
assert len(new_list) == 6, "the length of new_list must be 6"
```

##### 3.6 Extend ``new_list`` with ``tmp_list`` from above. Can you explain the difference between ``append`` and ``extend``?

<!-- grade_id: list_extend, points: 0, grade: false, solution: true -->
```python
new_list = my_list.copy()

# YOUR CODE HERE
raise NotImplementedError()

print(new_list)
print(f"length of new_list: {len(new_list)}")
```

<!-- grade_id: list_extend_test, points: 0.5, grade: true, solution: false -->
```python
# run tests to check your implementation
assert len(new_list) == 8, "the length of new_list must be 8"
```

##### 3.7 Create 2 variables
* ``var_1``: checks whether ``my_list`` contains "hello, world!"
* ``var_2``: checks whether ``my_list`` contains 4

<!-- grade_id: list_member, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

print(var_1)
print(var_2)
```

<!-- grade_id: list_member_test, points: 0.5, grade: true, solution: false -->
```python
# run tests to check your implementation
assert type(var_1) == bool
assert type(var_2) == bool
```

#### 4. if-statement
An if-statement is a control flow statement that allows you to execute a certain block of code only if a particular condition evaluates to true.
##### 4.1 Create an if - elif - else statement that checks whether a variable is equal to, greater, or less than 0. To do so:
* create a variable called ``x`` and assign it to -5
* if x is equal to 0: print "x is equal to 0"
* if x is less than 0: print "x is negative"
* if x is greater than 0: print "x is positive"
* play around with different values for x

<!-- grade_id: simple_if, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

##### 4.2 Implement a function that checks the potassium level of a person. The potassium level is usually between 3.7 and 5.3 mEq/l (milliequivalents per liter). To do so:
* create a function called ``check_potassium`` that takes ``value`` as an argument
* check if ``value`` is in the given normal potassium range
    * if yes: print an informative message (e.g. "the potassium level is in the correct range") and return True
    * if no: print an informative message and return False
* Play around with different values to see if your implementation works correctly

**Hint:** You can combine two conditions using logical operators. Also, you can create a variable to store your return value temporarily.

<!-- grade_id: check_potassium, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: check_potassium_test, points: 1, grade: true, solution: false -->
```python
# run tests to check your implementation
assert check_potassium(4.0) == True, "check if you used the correct potassium range"
```

##### 4.3 Implement a function that checks the hemoglobin level based on the age of the person. Assume the hemoglobin levels to be as follows:
* 0 - 18 years: 10 - 15.5 g/dl (grams per deciliter)
* \> 18 years: 12 - 18 g/dl (grams per deciliter)

To do so:
* create a function called ``check_hemoglobin`` that takes ``age`` and ``value`` as arguments
* create a nested if-construct
    * outer if-statement checks the range of ``age``
    * inner if-statements checks based on the age the correct range of hemoglobin level
        * if it is in the correct range print an informative message and return True
        * if it is **not** in the correct range print an informative message and return False



<!-- grade_id: check_hemoglobin, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: check_hemoglobin_test, points: 2, grade: true, solution: false -->
```python
# run tests to check your implementation
assert check_hemoglobin(16, 16.0) == False, "check if you used the correct hemoglobin range for underage people"
assert check_hemoglobin(23, 16.0) == True, "check if you used the correct hemoglobin range for adult people"
```

#### 5. Loops
Loops are a way to execute a block of code repeatedly. There are typically two types of loops:
1. For loops allow you to iterate over a sequence of elements (such as a list or range of numbers) and execute a block of code for each element in the sequence.
2. While loops allow you to execute a block of code repeatedly while a certain condition is true.
##### 5.1 Create a function ``calc_sum`` that takes a list of values as argument ``values``. Then calculate the sum of the list entries using a for loop:
**Hint:** Initialize a variable ``sum`` before the loop and assign it to 0. Test your solution with ``numbers``.

<!-- grade_id: for_sum, points: 0, grade: false, solution: true -->
```python
numbers = [23, 27, 21, 19]

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: for_sum_test, points: 1, grade: true, solution: false -->
```python
# run tests to check your implementation
numbers = [23, 27, 21, 19]
assert calc_sum(numbers) == 90, "check your for loop syntax and used operator"
```

##### 5.2 It is also possible to iterate over keys, values, and key-value pairs of a dictionary using a for loop. You can access them like in the following:
```
<dict>.keys() # get all the keys of dictionary <dict>
<dict>.values() # get all the values of dictionary <dict>
<dict>.items() # get all the key-value pairs of dictionary <dict>
```
Create a for loop that iterates over all keys of your dictionary ``patient`` and prints them. Then try the same for all values and items:

<!-- grade_id: for_dict, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

##### 5.3 Create a function ``arange`` which generates lists of ascending numbers using a while loop. It takes the following arguments:
* ``start``: start of the interval, included
* ``stop``: stop of the interval, excluded
* ``step``: spacing between values

To do so:
* create a control variable ``n`` and initialize it with ``start``
* initialize an empty list ``numbers``
* create the while loop such that it checks the current value of ``n`` and runs until ``stop``
* append ``n`` to ``list_numbers``
* add ``step`` in every iteration to your control variable.
* return ``numbers``

<!-- grade_id: while_arange, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: while_arange_test, points: 1, grade: true, solution: false -->
```python
# run tests to check your implementation
assert arange(5, 10, 1) == [5, 6, 7, 8, 9], "check if you set start and stop correctly, used the right operator, and performed all steps in the right order"
assert arange(5, 10, 2) == [5, 7, 9], "check if you set start and stop correctly, used the right operator, and performed all steps in the right order"
```

#### 6. Typecasting
Typecasting, also known as type conversion, is the process of converting a value of one data type to another data type. This can be useful in programming when you want to perform operations that require values of a specific data type.

##### 6.1 Convert the following variables to the specified data types:
* x from ``float`` to ``int``
* y from ``float`` to ``bool``
* z from ``int`` to ``float``

<!-- grade_id: typecasting, points: 0, grade: false, solution: true -->
```python
x = 8.7
y = 1.0
z = 9

# YOUR CODE HERE
raise NotImplementedError()

print(x)
print(y)
print(z)
```

<!-- grade_id: typecasting_test, points: 0.5, grade: true, solution: false -->
```python
# run tests to check your implementation
assert type(x) == int, "check if you converted x to the correct data type"
assert type(y) == bool, "check if you converted y to the correct data type"
assert type(z) == float, "check if you converted z to the correct data type"
```